```python
import os
```

```python
if os.path.exists('/coursedata'):
    COURSEDATA = '/coursedata'
else:
    COURSEDATA = '.'
```

```python
# I copied /usr/share/dict/words from my Linux workstation to the data directory
# using the workstation mount
```

```python
import pandas
```

```python
df = pandas.read_csv(os.path.join(COURSEDATA, 'words.txt'))
```

```python
df.head()
```

<!-- grade_id: cell-64c99be11824d217, points: 0, grade: false, solution: true -->
```python
# This is a student answer cell, so it must be marked "Autograded answer"

# Calculate the mean of all the length of all words
def mean_of_column(column):
    # YOUR CODE HERE
    raise NotImplementedError()
    return mean
```

```python
# This cell is read-only, because why not?

# Let's make and test some testdata
df_test = pandas.DataFrame({"testwords": ["dog", "cat", "camel"]})
assert abs(mean_of_column(df_test['testwords']) - 3.66666 ) < .0001
```

```python
# This is a normal cell, that students can edit.  We hope they play around
# with things here and learn something.
mean_of_column(df['A'])
```

<!-- grade_id: cell-b5655049d503321c, points: 0, grade: true, solution: false -->
```python
# This is marked "autograder tests", so can not be edited
assert abs(mean_of_column(df['A']) - 8.46449) < .0001
```

# Reflection
Please describe, in your own words, why the mean of the entire words file:

(this cell is "manually graded task" - markdown)

Please answer here

Write a function to test standard deviation of all words:

<!-- grade_id: cell-b6a45d75ebb21b41, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```